Two neuropeptides that promote blood feeding in Anopheles stephensi mosquitoes

Author Response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

Here Bansal et al., present a study on the fundamental blood and nectar feeding behaviors of the critical disease vector, Anopheles stephensi. The study encompasses not just the fundamental changes in blood feeding behaviors of the crucially understudied vector, but then use a transcriptomic approach to identify candidate neuromodulation path ways which influence blood feeding behavior in this mosquito species. The authors then provide evidence through RNAi knockdown of candidate pathways that the neuromodulators sNPF and Rya modulate feeding either via their physiological activity in the brain alone or through joint physiological activity along the brain-gut axis (but critically not the gut alone). Overall, I found this study to be built on tractable, well-designed behavioral experiments.

Their study begins with a well-structured experiment to assess how the feeding behaviors of A. stephensi changes over the course of its life history and in response to its age, mating and oviposition status. The authors are careful and validate their experimental paradigm in the more well-studied Ae. aegypti, and are able to recapitulate the results of prior studies which show that mating is pre-requisite for blood feeding behaviors in Ae. aegypt. Here they find A. stephensi like another Anopheline mosquitoes has a more nuanced regulation of its blood and nectar feeding behaviors.

The authors then go on to show in a Y- maze olfactometer that to some degree, changes in blood feeding status depend on behavioral modulation to host-cues, and this is not likely to be a simple change to the biting behaviors alone. I was especially struck by the swap in valence of the host-cues for the blood-fed and mated individuals which had not yet oviposited. This indicates that there is a change in behavior that is not simply desensitization to host-cues while navigating in flight, but something much more exciting happening.

The authors then use a transcriptomic approach to identify candidate genes in the blood feeding stages of the mosquito's life cycle to identify a list of 9 candidates which have a role in regulating the host-seeking status of A. stephensi. Then through investigations of gene knockdown of candidates they identify the dual action of RYa and sNPF and candidate neuromodulators of host-seeking in this species. Overrall, I found the experiments to be welldesigned. I found the molecular approach to be sound. While I do not think the molecular approach is necessarily an all-encompassing mechanism identification (owing mostly to the fact that genetic resources are not yet available in A. stephensi as they are in other dipteran models), I think it sets up a rich lines of research questions for the neurobiology of mosquito behavioral plasticity and comparative evolution of neuromodulator action.

Strengths:

I am especially impressed by the authors' attention to small details in the course of this article. As I read and evaluated this article I continued to think how many crucial details I may have missed if I were the scientist conducting these experiments. That attention to detail paid off in spades and allowed the authors to carefully tease apart molecular candidates of blood-seeking stages. The authors top down approach to identifying RYamide and sNPF starting from first principles behavioral experiments is especially comprehensive. The results from both the behavioral and molecular target studies will have broad implications for the vectorial capacity of this species and comparative evolution of neural circuit modulation.

I believe the authors have adequately addressed all of my concerns; however, I think an accompanying figure to match the explained methods of the tissue-specific knockdown would help readers. The methods are now explicitly written for the timing and concentrations required to achieve tissue-specific knockdown, but seeing the data as a supplement would be especially reassuring given the critical nature of tissue-specific knockdown to the final interpretations of this paper.

We thank the reviewer for the suggestion and have now incorporated a schematic in the supplementary figure S9B, explaining our methodology for achieving tissue-specific knockdowns.

Reviewer #2 (Public review):

Summary:

In this manuscript, Bansal et al examine and characterize feeding behaviour in Anopheles stephensi mosquitoes. While sharing some similarities to the well-studied Aedes aegypti mosquito, the authors demonstrate that mated-females, but not unmated (virgin) females, exhibit suppression in their blood-feeding behaviour. Using brain transcriptomic analysis comparing sugar fed, blood fed and starved mosquitoes, several candidate genes potentially responsible for influencing blood-feeding behaviour were identified, including two neuropeptides (short NPF and RYamide) that are known to modulate feeding behaviour in other mosquito species. Using molecular tools including in situ hybridization, the authors map the distribution of cells producing these neuropeptides in the nervous system and in the gut. Further, by implementing systemic RNA interference (RNAi), the study suggests that both neuropeptides appear to promote blood-feeding (but do not impact sugar feeding) although the impact was observed only after both neuropeptide genes underwent knockdown.

While the authors have addressed most of the concerns of the original manuscript, a few issues remain. Particularly, the following two points:

(5) Figure 4

The authors state that there is more efficient knockdown in the head of unfed females; however, this is not accurate since they only get knockdown in unfed animals, and no evidence of any knockdown in fed animals (panel D). This point should be revised in the results test as well.

Perhaps we do not understand the reviewer's point or there has been a misunderstanding. In Figure 4D, we show that while there is more robust gene knockdown in unfed females, bloodfed females also showed modest but measurable knockdowns ranging from 5-40% for RYamide and 2-21% for sNPF.

NEW-

In both the dsRNA treatments where animals were fed, neither was significantly different from control. Therefore, there is no change, and indeed this is confirmed by the author's labelling of the figure stats in panel 4D.

We agree with the reviewer and thank them for pointing it out. We have now revised the figure legend and the text to reflect these results (see lines 351-354).

In addition, do the uninjected and dsGFP-injected relative mRNA expression data reflect combined RYa and sNPF levels? Why is there no variation in these data,...

In these qPCRs, we calculated relative mRNA expression using the delta-delta Ct method (see line 975). For each neuropeptide its respective control was used. For simplicity, we combined the RYa and sNPF control data into a single representation. The value of this control is invariant because this method sets the control baseline to a value of 1.

NEW-

The authors are claiming that there is no variation between individual qPCR experiments (particularly in their controls)? Normally, one uses a known standard value (or calibrator) across multiple experiments/plates so that variation across biological replicates can be assessed. This has an impact on statistical analyses since there is no variation in the control data. Indeed, this impacts all figures/datasets in the manuscript where qPCR data is presented. All the controls have zero variation!

We are truly thankful to this reviewer for insisting on this point. It has made us revisit what we thought we understood and now realise were doing wrong (though many in literature do it this way!). We were – incorrectly – setting each control to 1 and calculating relative fold changes for each replicate independently. While this is often seen in literature, we now realise that it is incorrect. We have revisited all our analyses and normalized all samples to the mean ΔCt of the control group, which captures biological variation in both control and experimental groups. All data are now re-plotted to show individual data points for both control and experimental groups, and the error bars on controls represent the biological variation across replicates (Figure 4D, 4F, 4G, S8, S9). Statistical analyses were also revised accordingly, and, importantly, they do not change any conclusions. Please note that the abdominal expression of sNPF and RYa are so low that the controls show very variable baseline expression values.

Reviewer #3 (Public review):

Summary:

This manuscript investigates the regulation of host-seeking behavior in Anopheles stephensi females across different life stages and mating states. Through transcriptomic profiling, the authors identify differential gene expression between "blood-hungry" and "blood-sated" states. Two neuropeptides, sNPF and RYamide, are highlighted as potential mediators of host-seeking behavior. RNAi knockdown of these peptides alters host-seeking activity, and their expression is anatomically mapped in the mosquito brain (sNPF and RYamide) and midgut (sNPF only).

Strengths:

(1) The study addresses an important question in mosquito biology, with relevance to vector control and disease transmission.

Transcriptomic profiling is used to uncover gene expression changes linked to behavioral states.

(2) The identification of sNPF and RYamide as candidate regulators provides a clear focus for downstream mechanistic work.

(3) RNAi experiments demonstrate that these neuropeptides are necessary for normal hostseeking behavior.

(4) Anatomical localization of neuropeptide expression adds depth to the functional findings.

Weaknesses:

(1) The title implies that the neuropeptides promote host-seeking, but sufficiency is not demonstrated and some conclusions appear premature based on the current data. The support for this conclusion would be strengthened with functional validation using peptide injection or genetic manipulation.

(2) The identification of candidate receptors is promising, but the manuscript would be significantly strengthened by testing whether receptor knockdowns phenocopy peptide knockdowns. Without this, it is difficult to conclude that the identified receptors mediate the behavioral effects.

(3) Some important caveats, such as variation in knockdown efficiency and the possibility of offtarget effects, are not adequately discussed.

These comments were addressed in the previous round.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

Awesome paper everyone. A delight to read and review.

Thank you very much! We appreciated your comments too!

Changes from the earlier version to the revised version

Abstract

Animals routinely need to make decisions about what to eat and when. These decisions are influenced not only by the availability and quality of food but also by the internal state of the animal, which needs to compute and give weights to these different variables before making a choice. Feeding preferences of female mosquitoes exemplify this behavioural plasticity. Both male and female mosquitoes usually feed on carbohydrate-rich sources of nectar or sap, but the female also feeds on blood, which is essential for egg development. This blood-appetite is modulated across the female’s reproductive cycle, yet little is known about the factors that bring it about. We show that mated, but not virgin Anopheles stephensi females, a major vector of urban malaria in the Indian sub-continent and West Africa, suppress blood feeding between a blood meal and oviposition. We identify several candidate genes through transcriptomics of blood-deprived and -sated An. stephensi central brains that could modulate this behaviour. We show that short neuropeptide F (sNPF) and RYamide (RYa) act together to promote blood feeding and identify a cluster of cells in the subesophageal zone that expresses sNPF transcripts only in the blood-hungry state. Such females also have more sNPF transcripts in their midguts. Based on these data, we propose a model where increased sNPF levels in the brain and gut promote a state of blood-hunger, which drives feeding behaviour either by sNPF’s action in the two tissues independently or via a communication between them. This occurs in the context of the action of RYa in the brain.

Introduction and Results

An animal’s dietary choice is an outcome of its internal physiological state and external stimuli. Its changing nutritional demands, reproductive state, and developmental history are integrated with external cues before a food choice is made. For example, when deprived of a particular nutrient, like proteins, animals show a specific increase in protein-appetite (1–3). Mating influences food choices of female flies by increasing their preference for proteins, sugars, and salt for egg development and oviposition (4–8).

Changes in both sensory perception and central brain circuits have been shown to underlie such behavioural changes. For example, sensory neurons of starved animals become more sensitive to attractive cues and less sensitive to aversive ones, both in vertebrates and invertebrates (9–16). These changes are often brought about by neuromodulators and neuropeptides (9, 10, 17–22). In the central brain, internal nutrient or hormone sensors can monitor the animal’s nutritional state. In mice, neurons in the hypothalamus express receptors for hunger and satiety signals like ghrelin and leptin (16, 23, 24). In insects, several neurons in the mushroom body, central complex, and subesophageal zone (SEZ) have been shown to integrate hunger and satiety signals to regulate feeding on sugars or proteins, choose between exploration or exploitation of food sources, or initiate or terminate feeding (25–33).

Female mosquitoes offer a striking model for dissecting how internal state and external cues shape dietary choices. Many female mosquitoes have a dual appetite for carbohydrate-rich nectar or sap – to meet energetic requirements – and protein-rich blood – for egg development. Which meal a female chooses to take is influenced by nutritional demands, reproductive state, and developmental history. For example, mosquitoes with low teneral reserves will take multiple blood meals for egg development (34–37). Nutritional state also influences other internal state outputs on blood feeding. For example, whether virgin females take blood meals or not is dependent on prior sugar feeding in most Aedes strains ((38) and references within).

Despite these species- and context-dependent variations, a consistent finding is that female mosquitoes oscillate between states of heightened and suppressed host seeking across successive reproductive cycles (figure 1A). For example, in Ae. aegypti and An. gambiae, newly emerged mosquitoes display no or low interest in the host, which increases as they age (34, 39–41) and once fed to repletion, host attraction is suppressed. The strength and timing of this suppression varies between species — being robust and well-characterised in mated Ae. aegypti (42–44), and more flexible in anophelines (36, 45, 46). Sugar feeding is not modulated in this way, in some cases it is preferred soon after emergence (40, 47).

Feeding behaviours of An. stephensi.

A. Reproductive cycle of An. stephensi. Upon emergence [1, D0-D5] females have a dual appetite for protein-rich blood and carbohydrate-rich sources of sugar [2]. After a blood meal [3,4,5, D1PBM-D4PBM], the eggs mature and are laid in water [6]. The next reproductive cycle can now begin. Striped abdomens indicate mated females. B. Blood-feeding assay. 0-7h-old mosquitoes were collected in cups and aged appropriately on sugar. Prior to the test, mosquitoes were ‘activated’ for 5 minutes by presenting a human hand, followed by 3 exhalations. Blood meals were presented through a Hemotek perfumed with human skin odours. Blood-fed mosquitoes were visually scored. C. Blood-feeding behaviour of Ae. aegypti. 8 days-old virgin and mated females were assessed for first blood meal. Fully fed mated females were assessed for second blood meal (D1PBM-D4PBM). 18-21 females/trial, n=10-12 trials/group. Kruskal-Wallis with Dunn’s multiple comparisons test, p<0.05. Data labelled with different letters are significantly different. (m): mated. D. Blood-feeding behaviour of An. stephensi. Co-housed and virgin females were tested on the day of emergence (D0) and each day for the next 5 days (D1-D5). D1PBM-D4PBM represents blood feeding 1-4 days post-first blood meal and D1PO-D2PO represents blood feeding 1-2 days post-oviposition. 18-21 females/trial, n=9-20 trials/group. Generalized Linear Mixed Model with post-hoc pairwise comparisons using estimated marginal means and Bonferroni-corrected p-values; *p<0.05; **p<0.01; ****p<0.0001. E. Sugar-feeding behaviour of An. stephensi. Females of similar conditions as in (D) were tested for sugar feeding. 18-21 females/trial, n=9-20 trials/group. Generalized Linear Mixed Model with post-hoc pairwise comparisons using estimated marginal means and Bonferroni-corrected p-values; **p<0.01; ***p<0.001; ****p<0.0001. F. Choice assay for blood and sugar. Sugar-sated or sugar-starved females, co-housed with males (indicated by striped abdomens) were given a choice between blood and sugar and assessed for the choice made (bottom): blood and sugar (blue); blood only (magenta); sugar only (orange); none (grey). Co-housed males were assessed for sugar feeding only. Proportion of mosquitoes fed on each particular choice are shown at the bottom. 17-24 mosquitoes/trial; n=11-12 trials/condition. One-way ANOVA with Holm-Šídák multiple comparisons test, p<0.05. Data labelled with different letters are significantly different. G. Mating and blood feeding. Mating status of D0-D5 co-housed females assayed in (D) were determined post-hoc via spermatheca dissection. n=232-239 females analysed for each time point. H. Mating after blood feeding. Blood-fed virgins were allowed to mate and assayed for second blood meal. Age-matched virgins were used as controls. 15-20 females/trial, n=10 trials. Unpaired t-test; ****p<0.0001.

The modulation of blood feeding provides a useful framework for understanding how internal and external cues shape meal choice. What underlying molecular cues represent internal physiological states to influence dietary choice is of interest not only because of its importance in public health, but also because such understanding has implications for other organisms. In Ae. albopictus anticipatory expression of the yolk protein precursor gene vitellogenin 2 in the fat body promotes host-seeking behaviour (48). In Ae. aegypti, neuropeptide Y-like receptor 7 (NPYLR7), and more recently, short neuropeptide F (sNPF) and neuropeptide RYamide (RYa) have been shown to suppress host attraction, while neuropeptide F (NPF) promotes it (44, 49–52). Factors from the host can also both promote (ATP) (53) and terminate feeding (Fibrinopeptide A)(54).

Here we describe the blood-and sugar-feeding patterns of Anopheles stephensi, which is endemic to the Indian sub-continent but has expanded its range into West Africa, putting millions more people at risk of malaria (55). Neurotranscriptomics across the different conditions of blood-deprivation and -satiety allowed us to shortlist several candidate molecules that might promote blood-feeding behaviour. Through dsRNA-mediated knockdown of nine of them, we identified two neuropeptide genes – sNPF and RYa - that together promote blood feeding through their action in the brain and possibly the abdomen. We find that sNPF expression pattern reflects this requirement: a cluster of cells in the SEZ expresses sNPF transcripts only in the blood-hungry females and such females also have more sNPF transcripts in their midguts. RYa is not modulated similarly. Both receptors are expressed in the brain and the sNPF receptor is also expressed in the midgut. Together, these data are consistent with a model where an increase in levels of sNPF in the brain and the midgut promotes a state of blood-hunger in An. stephensi. This occurs in the context of RYa signalling in the brain.

Blood feeding, but not sugar feeding, is modulated across the reproductive cycle

Earlier version: To assess the female An. stephensi appetite for blood through the reproductive cycle we designed a behavioural assay (figure 1B, S1A; see methods for details), and tested it on Ae. aegypti, a well-studied model. We were able to reproduce the previously reported results that mating enhances blood feeding in Aedes (56–58) and that once fed to repletion, females actively suppress blood feeding until oviposition (59) (figure 1C). We used this assay on An. stephensi, and employed Generalised Linear Mixed Model (GLMM) to compare appetites across conditions (see supplementary sheet 1). We found that these females were uninterested in blood when they emerge but that their appetite increased with age (magenta, D0-D5 in figure 1D, supplementary sheet 1:model 1). Once fed to repletion, like Ae. aegypti, An. stephensi also suppress blood feeding for at least four days after a blood meal (magenta, D1PBM-D4PBM in figure 1D) and until oviposition, after which, the suppression was lifted to pre-blood meal levels (magenta, D1PO-D2PO in figure 1D). Unlike blood feeding, sugar feeding was not modulated similarly barring a brief suppression post-blood meal, which was not dependent on oviposition (figure 1E; methods and S1G for details, supplementary sheet 1:model 4, 5, 6).

Revised version: To assess the female An. stephensi appetite for blood through the reproductive cycle we designed a behavioural assay (figure 1B, S1A; see methods for details), and tested it on Ae. aegypti, a well-studied model. We were able to reproduce the previously reported results that mating enhances blood feeding in Aedes (56– 58) and that once fed to repletion, females actively suppress blood feeding until oviposition (59) (figure 1C).

Added in the revised version: We used this assay on An. stephensi, and employed Generalised Linear Mixed Model (GLMM) to compare appetites across conditions (see supplementary sheet 1). We found that these females were uninterested in blood when they emerge but that their appetite increased with age (magenta, D0-D5 in figure 1D, supplementary sheet 1: model 1). Once fed to repletion, like Ae. aegypti, An. stephensi also suppress blood feeding for at least four days after a blood meal (magenta, D1PBM-D4PBM in figure 1D) and until oviposition, after which, the suppression was lifted to pre-blood meal levels (magenta, D1PO-D2PO in figure 1D). Unlike blood feeding, sugar feeding was not modulated similarly barring a brief suppression post-blood meal, which was not dependent on oviposition (figure 1E; methods and S1G for details, supplementary sheet 1: model 4, 5, 6).

Earlier version: To ascertain whether sugar and blood represent parallel appetites independent of each other, we presented sugar-sated co-housed females a choice between the two. Under these conditions, 70% of the females chose to feed on blood alone and only 21% took both meals (figure 1F). When co-housed females were deprived of sugars for a day prior to testing, the pattern reversed: 68% of the females now took both sugar and blood meals and only 24% fed on blood alone (figure 1F). As expected, males fed robustly on sugar in both condi-tions (grey, figure 1F).

Revised version: To ascertain whether sugar and blood represent parallel appetites independent of each other, we presented sugar-sated co-housed females a choice between the two. Under these conditions, 70% of the females chose to feed on blood alone and only 21% took both meals (figure 1F). When co-housed females were deprived of sugars for a day prior to testing, the pattern reversed: 68% of the females now took both sugar and blood meals and only 24% fed on blood alone (figure 1F). As expected, males fed robustly on sugar in both conditions (grey, figure 1F).

Since females took blood meals regardless of their prior sugar feeding status and only sugar feeding was selectively suppressed by prior sugar access, we infer that these two represent parallel appetites regulated independently. The fact that most females in the starved group took both sugar- and blood meals, while most females in the sated condition ignored sugar, supports this interpretation. An alternate interpretation of these data are that two represent hierarchical appetites: that some sugar feeding is necessary to stimulate blood feeding. Determining the order of feeding events in the choice assay would be necessary to comment on this. However, prior reports that females can survive on blood meals alone (with no access to plant sugars (36)) argue against this model. Together, these data suggest that in newly emerged females, blood and sugar represent parallel appetites. We note that their interdependence may emerge later in the reproductive cycle (D1PBM-D2PBM in figure 1E).

Earlier version: As mating promotes the first blood meal in Ae. aegypti, we tested if this was true for An. stephensi. A post-hoc dissection of each female’s spermatheca (where sperm is stored after mating) revealed that many of the blood-fed females in the above assay were unmated (figure 1G). To confirm this, we assayed blood-feeding behaviour in virgin females and found that virgin An. stephensi, unlike Ae. aegypti, have a robust appetite for blood and that mating enhances it only marginally (cyan, D0-D5, figure 1D, supplementary sheet 1:model 1, 2). Interestingly, however, these virgins, unlike mated An. stephensi (or Ae. aegypti) females, failed to suppress their blood-appetite, even after engorging on blood. Instead, they continued to blood-feed for up to four days after their first blood meal (cyan, D1PBM - D4PBM in figure 1D, supplementary sheet 1:model 3). A flipped order of the two behaviours - virgins were first given a blood meal then allowed to mate – also resulted in suppression of the appetite (figure 1H). This suggests that the post-blood meal suppression of blood-appetite in An. stephensi is dependent upon mating status.

Revised version: As mating promotes the first blood meal in Ae. aegypti, we tested if this was true for An. stephensi. A post-hoc dis-section of each female’s spermatheca (where sperm is stored after mating) revealed that many of the blood-fed females in the above assay were unmated (figure 1G). To confirm this, we assayed blood-feeding behaviour in virgin females and found that virgin An. stephensi, unlike Ae. aegypti, have a robust appetite for blood and that mating enhances it only marginally (cyan, D0-D5, figure 1D, supplementary sheet 1: model 1, 2). Interestingly, however, these virgins, unlike mated An. stephensi (or Ae. aegypti) females, failed to suppress their blood-appetite, even after engorging on blood. Instead, they continued to blood-feed for up to four days after their first blood meal (cyan, D1PBM - D4PBM in figure 1D, supplementary sheet 1: model 3). A flipped order of the two behaviours - virgins were first given a blood meal then allowed to mate – also resulted in suppression of the appetite (figure 1H). This suggests that the post-blood meal suppression of blood-appetite in An. stephensi is dependent upon mating status.

Together, these data describe female An. stephensi’s feeding behaviours: 1) in newly emerged females blood and sugar represent parallel appetites, 2) blood-, but not sugar-appetite (barring a brief effect), is suppressed between a blood meal and oviposition, and 3) mating is not necessary for blood feeding in newly emerged females, but it is necessary for its suppression between a blood meal and oviposition.

Host seeking is modulated across the reproductive cycle

The modulation of blood feeding through the reproductive cycle could be due to 1) modulation in peripheral sensitivity to host kairomones, 2) modulation by central brain circuits, or 3) modulation in both peripheral and central circuit elements. Extensive data in other mosquito species suggest that sensitivity to host kairomones change as females age (39, 60–63) and following a blood meal (46, 64, 65). We tested this in An. stephensi.

Earlier version: Because of the multi-sensory nature of host seeking, we used a two-choice behavioural assay where multiple host cues were provided in one arm while the other had none. Mosquitoes had to make a choice between the two arms (figure 2A). When no host cues are provided in either (blank) even blood-hungry females distribute equally between the two arms demonstrating the assay has no bias (black, figure 2B). Using this assay, we found that newly emerged females showed no preference for either the control or the host arm (figure 2B). Five days later, they were strongly attracted to the host cues and this attraction was suppressed after their blood meal and until oviposition (magenta, figure 2B). This general trend in attraction to host cues tracks the trend in blood feeding. Interestingly, the response of virgin females to host cues in the Y-maze was highly variable under all conditions tested (cyan, figure 2B). Yet, like in the blood-feeding assay, bloodfed virgins were comparatively more attracted to human host cues than were the mated blood-fed females (figure 2B).

Revised version: Because of the multi-sensory nature of host seeking, we used a two-choice behavioural assay where multiple host cues were provided in one arm while the other had none. Mosquitoes had to make a choice between the two arms (figure 2A). When no host cues are provided in either (blank) even blood-hungry females distribute equally between the two arms demonstrating the assay has no bias (black, figure 2B). Using this assay, we found that newly emerged females showed no preference for either the control or the host arm (figure 2B). Five days later, they were strongly attracted to the host cues and this attraction was suppressed after their blood meal and until oviposition (magenta, figure 2B). This general trend in attraction to host cues tracks the trend in blood feeding. Interestingly, the response of virgin females to host cues in the Y-maze was highly variable under all conditions tested (cyan, figure 2B). Yet, like in the blood-feeding assay, blood-fed virgins were comparatively more attracted to human host cues than were the mated blood-fed females (figure 2B).

Host seeking is modulated like blood feeding is in An. stephensi.

A. Y-maze schematic for host-seeking behaviour of female An. stephensi. Females were acclimatised in the chamber for 5 minutes while being exposed to host kairomones presented in a test arm. A fan sucked the air from both host and control arms at 0.3-0.6m/s. Post-acclimatisation, mosquitoes were released and allowed to choose between the two arms. B. Percent females attracted to either host cues or control arm at the indicated time points. 17-21 females/trial, n=10 trials/group. Unpaired t-test; n.s.: not significant, p>0.05.; * p<0.05; ****p<0.0001. (D0: day of emergence; D5(m): 5 days post-emergence, mated; D5: 5 days post-emergence, virgin; D3PBM(m): 3 days post-blood meal, mated; D3PBM: 3 days post-blood meal, virgin; D2PO(m): 2 days post-oviposition, mated)

These data suggest that, like blood feeding, host seeking is also modulated through the female’s reproductive cycle, and mating influences this modulation by enhancing it initially and suppressing it post-blood meal. The variability in response of virgin females suggests that mating enhances the female’s initial attraction to a host.

Neurotranscriptomics reveals several candidate genes that promote blood feeding

While sensitivity to host kairomones may modulate blood feeding, it is also possible that central brain circuitry contributes to it. We reasoned that cues from mating, or signals from other tissues such as the midgut could act on brain circuitry to modulate blood feeding. To determine the molecular underpinnings of this behavioural modulation, we performed RNA sequencing from the central brains (brains without optic lobes) of mosquitoes at different stages of their reproductive cycle (consequently different states of blood-deprivation or -satiety). This included females at emergence (D0), 5-days after emergence (D5; sugar-fed, blood-deprived), virgins and mated females one day after a blood meal (D1PBM, D1PBM(m)), and females one day after oviposition (D1PO(m)). Age-matched males were included for the earliest two time points (figure 3A).

Neurotranscriptome of An. stephensi.

Earlier version: A. Central brains sampled from males and females at different ages and feeding preferences for bulk RNA-seq. Uninterested in blood: D0 males, D0 females, D5 males. Blood-hungry: D5, sugar-fed virgin females, D1PBM virgin females, and D1PO(m) mated females. Blood-sated: D1PBM(m) mated females. n=100 central brains per replicate; 3 biological replicates per sample, 7 samples. (m): mated. B. Principal Component Analysis (PCA) of the normalised counts-permillion. C. Genes (listed on the right) with known expression patterns plotted across different samples. Average of the triplicate data are represented as z-score normalised log10 (TPM+1) values. Grey cells represent no data. D. Identification of candidates potentially involved in promoting blood-feeding behaviour. Two types of comparisons were made: (left) genes commonly upregulated in the three blood-hungry conditions (see figure 3A) when compared against D0 females and excluded in males, and (right) genes upregulated in D5 sugar-fed virgin females when compared against D1PBM mated females (blood-sated), and excluded in males. Number of candidate genes identified from two different sets of analyses are boxed in red. E. Nine candidate genes (listed on the right), shortlisted for further validation. RNA-seq data from central brain samples in triplicates are represented as z-score normalised log10 (TPM+1) values. (D0: day of emergence; D5: 5 days post-emergence, virgin; D1PBM: 1 day post-blood meal, virgin; D1PBM(m): 1 day post-blood meal, mated; D1PO(m): 1 day post-oviposition, mated)

Revised version: A. Central brains sampled from males and females at different ages and feeding preferences for bulk RNA-seq. Uninterested in blood: D0 males, D0 females, D5 males. Blood-hungry: D5, sugar-fed virgin females, D1PBM virgin females, and D1PO(m) mated females. Blood-sated: D1PBM(m) mated females. n=100 central brains per replicate; 3 biological replicates per sample, 7 samples. (m): mated. B. Principal Component Analysis (PCA) of the normalised counts-per-million. C. Genes (listed on the right) with known expression patterns plotted across different samples. Average of the triplicate data are represented as z-score normalised log10 (TPM+1) values. Grey cells represent no data. D. Identification of candidates potentially involved in promoting blood-feeding behaviour. Two types of comparisons were made: (left) genes commonly upregulated in the three blood-hungry conditions (see figure 3A) when compared against D0 females and excluded in males, and (right) genes upregulated in D5 sugar-fed virgin females when compared against D1PBM mated females (blood-sated), and excluded in males. Number of candidate genes identified from two different sets of analyses are boxed in red. E. Nine candidate genes (listed on the right), shortlisted for further validation. RNA-seq data from central brain samples in triplicates are represented as z-score normalised log10 (TPM+1) values. (D0: day of emergence; D5: 5 days post-emergence, virgin; D1PBM: 1 day post-blood meal, virgin; D1PBM(m): 1 day post-blood meal, mated; D1PO(m): 1 day post-oviposition, mated)

Earlier version: For each of our samples, 81-87% of about 20-30 million reads aligned to the An. stephensi genome (66). The variance between our replicates was low as represented in the Principal Component Analysis (PCA) (figure 3B). In PCA, older female brains, regardless of mating-, blood feeding-, or ovipositionstatus clustered together while the youngest female and male brains stood apart on PC1 and PC2 suggesting that age and sex contribute most significantly to the variance in our data (figure 3B). The expression pattern of some hallmark genes showed the expected trends. Genes involved in eclosion behaviours – eclosion hormone (67) and partner of bursicon (68) – showed the highest expression in the D0 samples and were not expressed later. Male-specific genes, HMG-B-13-like, were expressed only in males (69), and genes involved in vitellogenic stages of egg-development – vitellogenins, ecdysone receptor, and E-20-monooxygenase (70)– were expressed only in post-blood meal samples (figure 3C).

Revised version: For each of our samples, 81-87% of about 20-30 million reads aligned to the An. stephensi genome (66). The variance between our replicates was low as represented in the Principal Component Analysis (PCA) (figure 3B). In PCA, older female brains, regardless of mating-, blood feeding-, or oviposition-status clustered together while the youngest female and male brains stood apart on PC1 and PC2 suggesting that age and sex contribute most significantly to the variance in our data (figure 3B). The expression pattern of some hallmark genes showed the expected trends. Genes involved in eclosion behaviours – eclosion hormone (67) and partner of bursicon (68) – showed the highest expression in the D0 samples and were not expressed later. Male-specific genes, HMG-B-13-like, were expressed only in males (69), and genes involved in vitellogenic stages of egg-development – vitellogenins, ecdysone receptor, and E-20-monooxygenase (70)– were expressed only in post-blood meal samples (figure 3C).

We next explored this data for insights into the molecular underpinnings of blood-feeding modulation. Reports suggest that anticipatory metabolic changes in female germline cells and a ‘vittelogenin wave’ in fat bodies can drive specific appetite in Drosophila (71) and Ae. albopictus (48), respectively. However, we did not find any evidence in the neurotranscriptomes for a similar anticipatory metabolic change that might drive blood feeding in the mosquito brain (figure S2B). Instead, our data suggests altered carbohydrate metabolism after a blood meal, with the female brain potentially entering a state of metabolic ‘sugar rest’ while actively processing proteins (Figure S2B, S3). However, physiological measurements of carbohydrate and protein metabolism will be required to confirm whether glucose is indeed neither spent nor stored during this period. We similarly plotted genes involved in lipid metabolism, neurotransmitters, neuropeptides and their receptors and found no significant patterns (figure S4, S5, S6). In summary, contrary to our expectations, we saw no evidence in the neurotranscriptomes for anticipatory metabolic changes in the brain that might in turn drive blood-feeding behaviour. It is possible, however, that other forms of regulation that do not involve transcriptional changes could still influence metabolism.

In an alternate approach to identifying genes that might modulate blood feeding, we took the intersection of genes that were enriched in the three blood-hungry conditions (and absent in males). These comparisons were made against the two conditions that were not motivated to blood-feed: females at emergence and mated blood-fed females (figure 3D). These shortlisted 38 and 34 genes respectively, gave us 68 unique candidates (excluding common genes and isoforms) that might promote blood feeding in An. stephensi (figure S7). Notably, we did not find any genes to be differentially expressed between the two blood-fed conditions (virgins and mated), despite their contrasting behaviours (figure S2A). The strongest signatures observed in these females were those related to metabolism (figure S2B). It is possible that the differences lie in small populations of cells, which get obscured in the averaged transcriptomes of the entire brain.

Neuropeptides sNPF and RYamide promote blood feeding in An. Stephensi

Earlier version: We employed multiple, independent criteria while further shortlisting candidates from the above analysis: 1) genes that showed the expected expression pattern: low in newly emerged males and females; high in blood-hungry females; low in blood-sated females, 2) all neuropeptides, since these are known to regulate feeding behaviours, and 3) all genes related to juvenile hormone (JH) and 20-hydroxyecdysone (20E) signalling pathway, since they regulate development and reproduction in mosquitoes. This gave us 9 candidates (figure 3E): Two sodium-dependent amino acid transporters 118504075 and 118504077. Transporters can sense nutrients and regulate feeding in Drosophila (29, 72). The neuropeptides included Orcokinin (Ork), Prothoracicostatic peptides (Prp), RYamide (RYa) and short neuropeptide F (sNPF), and a metabolic hormone Eiger, which responds to low nutrient conditions in Drosophila (73). Protein Hairy and Juvenile hormone epoxide hydrolase 1-like (JHEH) are involved in the JH pathway (74, 75)(figure 3E). Using dsRNA, we knocked each of these genes down and tested their blood-feeding behaviour. In each case, we ensured efficiency of gene knockdown post-hoc (figure 4A).

Revised version: We employed multiple, independent criteria while further shortlisting candidates from the above analysis: 1) genes that showed the expected expression pattern: low in newly emerged males and females; high in blood-hungry females; low in blood-sated females, 2) all neuropeptides, since these are known to regulate feeding behaviours, and 3) all genes related to juvenile hormone (JH) and 20-hydroxyecdysone (20E) signalling pathway, since they regulate development and reproduction in mosquitoes. This gave us 9 candidates (figure 3E): Two sodium-dependent amino acid transporters 118504075 and 118504077. Transporters can sense nutrients and regulate feeding in Drosophila (29, 72). The neuropeptides included Orcokinin (Ork), Prothoracicostatic peptides (Prp), RYamide (RYa) and short neuropeptide F (sNPF), and a metabolic hormone Eiger, which responds to low nutrient conditions in Drosophila (73). Protein Hairy and Juvenile hormone epoxide hydrolase 1-like (JHEH) are involved in the JH pathway (74, 75)(figure 3E). Using dsRNA, we knocked each of these genes down and tested their blood-feeding behaviour. In each case, we ensured efficiency of gene knockdown post-hoc (figure 4A). We first confirmed that the injection itself did not cause any behavioural changes (figure 4B) and in all following experiments either dsGFP-injected or uninjected females were included as controls.

short neuropeptide F (sNPF) and RYamide (RYa) together promote blood feeding in An. stephensi.

Earlier version: A. Pipeline used for functional validation: dsRNAs were injected in adult virgins and tested for blood feeding. Knockdown efficiency was determined in heads and abdomens of fed and/or unfed females after the behaviour. B. Blood feeding in uninjected and dsGFP-injected females. 19-25 females/replicate, n=8-9 replicates/group. Unpaired t-test; n.s.: not significant, p>0.05. C. Blood- and sugar-feeding behaviour of females where both RYa and sNPF were knocked down in both the heads and abdomens. Controls: Uninjected females. 14-25 females/replicate, n=8-9 replicates/group. Unpaired t-test; n.s.: not significant, p>0.05; ****p<0.0001. D. Relative mRNA expressions of RYa and sNPF in the heads of dsRYa+dssNPF - injected blood-fed and unfed females assayed in (C). mRNA levels were compared to those in the uninjected controls. n=6 replicates/group. One-way ANOVA with Holm-Šídák multiple comparisons test, p<0.05. Data labelled with different letters are significantly different. E. Blood- and sugar-feeding behaviour of females when both RYa and sNPF were knocked down only in the abdomens. Controls: dsGFP-injected females. 17-21 females/replicate, n=3 replicates/group. Unpaired t-test; n.s.: not significant, p>0.05. F,G. Relative mRNA expressions of RYa and sNPF in the heads (F) and abdomens (G) of dsRYa+dssNPF - injected fed females assayed in (E). mRNA levels were compared to those in the dsGFP-injected controls. n=3 replicates/group. One-way ANOVA with Holm-Šídák multiple comparisons test, p<0.05. Data labelled with different letters are significantly different. H. RYa (cyan) and sNPF (magenta) HCR in situ hybridisation in central brain of 5 days-old sugar-fed virgin female. nc82 (blue). Scale bar, 50µm. I. Co-expression of RYa (cyan) and sNPF (magenta) in mushroom body Kenyon cells of 5 days-old sugar-fed virgin female. Scale bar, 50µm. J. sNPF expression in the SEZ of females uninterested in blood (D0), blood-hungry females (D5) and blood-sated females (D2PBM(m)). Magenta arrow marks the novel sNPF cluster only in the blood-hungry (D5) condition. Scale bar, 50µm. K. sNPF HCR in situ hybridisation in gut of 5 days-old sugar-fed virgin female. Higher magnification image in shown in K’. Scale bar, 50µm. L. sNPF expression in the posterior midguts of females uninterested in blood (D0), blood-hungry females (D5) and blood-sated females (D2PBM(m)). Note expression in enteroendocrine cells (K’) and in two cells in the anterior midgut (L, insets). M. No. of sNPF positive cells. n=5-7 guts per condition. One-way ANOVA with Holm-Šídák multiple comparisons test, p<0.05. Data labelled with different letters are significantly different. N. sNPF receptor (sNPFR) and RYa receptor(RYaR) HCR in situ hybridisation of central brain (left) and midgut (right) of 5 days-old sugar-fed virgin female. Scale bar, 50µm. O. Proposed model: RYa and RYaR (cyan) are expressed only in the brain, while sNPF and sNPFR (magenta) are expressed both in the brain and the gut. Increased sNPF levels in both the tissues (dashed circles) promote a state of blood-hunger, which may drive feeding behaviour either by sNPF’s action in the two tissues independently or via a communication between them. This happens in the context of RYa signalling in the brain. (D0: day of emergence; D5: 5 days post-emergence, virgin; D2PBM(m): 2 days post-blood meal, mated)

Revised version: A. Pipeline used for functional validation: dsRNAs were injected in adult virgins and tested for feeding behaviours. Knockdown efficiency was determined in heads and abdomens of fed and/or unfed females after the behaviour. B. Blood feeding in uninjected and dsGFP-injected females. 19-25 females/replicate, n=8-9 replicates/group. Unpaired t-test; n.s.: not significant, p>0.05. C. Blood- and sugar-feeding behaviour of females where both RYa and sNPF were knocked down in both the heads and abdomens. Controls: Uninjected females. 14-25 females/replicate, n=8-9 replicates/group. Unpaired t-test; n.s.: not significant, p>0.05; ****p<0.0001. D. Relative mRNA expressions of RYa and sNPF in the heads of dsRYa+dssNPF - injected blood-fed and unfed females assayed in (C). mRNA levels were compared to those in the uninjected controls. n=5 replicates/group. One-way ANOVA with Holm-Šídák multiple comparisons test for relative RYa expression and Kruskal-Wallis with Dunn’s multiple comparisons test for relative sNPF expression, p<0.05. Data labelled with different letters are significantly different. E. Blood- and sugar-feeding behaviour of females when both RYa and sNPF were knocked down only in the abdomens. Controls: dsGFP-injected or uninjected females. 17-22 females/replicate, n=5 replicates/group. Unpaired t-test; n.s.: not significant, p>0.05. F,G. Relative mRNA expressions of RYa (F) and sNPF (G) in the heads and abdomens of dsRYa+dssNPF - injected fed females assayed in (E). mRNA levels in both tissues were compared to their respective dsGFP-injected or uninjected females controls. n=5 replicates/group. Unpaired t-test; n.s.: not significant, p>0.05; **p<0.01; ****p<0.0001. H. RYa (cyan) and sNPF (magenta) HCR in situ hybridisation in central brain of 5 days-old sugar-fed virgin female. nc82 (blue). Scale bar, 50µm. I. Co-expression of RYa (cyan) and sNPF (magenta) in mushroom body Kenyon cells of 5 days-old sugar-fed virgin female. Scale bar, 50µm. J. sNPF expression in the SEZ of females uninterested in blood (D0), blood-hungry females (D5) and blood-sated females (D2PBM(m)). Magenta arrow marks the novel sNPF cluster only in the blood-hungry (D5) condition. Scale bar, 50µm. K. sNPF HCR in situ hybridisation in gut of 5 days-old sugar-fed virgin female. Higher magnification image in shown in K’. Scale bar, 50µm. L. sNPF expression in the posterior midguts of females uninterested in blood (D0), blood-hungry females (D5) and blood-sated females (D2PBM(m)). Note expression in enteroendocrine cells (K’) and in two cells in the anterior midgut (L, insets). M. No. of sNPF positive cells. n=5-7 guts per condition. One-way ANOVA with Holm-Šídák multiple comparisons test, p<0.05. Data labelled with different letters are significantly different. N. sNPF receptor (sNPFR) and RYa receptor(RYaR) HCR in situ hybridisation of central brain (left) and midgut (right) of 5 days-old sugar-fed virgin female. Scale bar, 50µm. O. Proposed model: RYa and RYaR (cyan) are expressed only in the brain, while sNPF and sNPFR (magenta) are expressed both in the brain and the gut. Increased sNPF levels in both the tissues (dashed circles) promote a state of blood-hunger, which may drive feeding behaviour either by sNPF’s action in the two tissues independently or via a communication between them. This happens in the context of RYa signalling in the brain. (D0: day of emergence; D5: 5 days post-emergence, virgin; D2PBM(m): 2 days post-blood meal, mated)

Removed in the revised version: We first confirmed that the injection itself did not cause any behavioural changes (figure 4B) and in all following experiments either dsGFP-injected or uninjected females were included as controls.

Earlier version: We saw no change in behaviour upon loss of Eiger, JHEH, and Ork function (figure S8A-C, S8A’-C’). We were able to downregulate Hairy and Prp transcript levels by only 30-35%, hence cannot comment on their role in modulating blood feeding (figure S8D, S8E, S8D’, S8E’). Downregulation of the two transporters was efficient and resulted in a marginal reduction in blood feeding (figure S8F, S8G, S8F’, S8G’). However, their combined knockdowns, though effective, did not show any significant change in behaviour (figure S8H, S8H’).

Revised version: We saw no change in behaviour upon loss of Eiger, JHEH, and Ork function (figure S8A-C, S8A’-C’). We were able to downregulate Hairy and Prp transcript levels by only 25-38%, hence cannot comment on their role in modulating blood feeding (figure S8D, S8E, S8D’, S8E’). Downregulation of the two transporters was efficient and resulted in a marginal reduction in blood feeding (figure S8F, S8G, S8F’, S8G’). However, their combined knockdowns, though effective, did not show any significant change in behaviour (figure S8H, S8H’).

Earlier version: Two neuropeptides - sNPF and RYa - showed about 25% and 40% reduced mRNA levels in the heads but the proportion of females that took blood meals remained unchanged (figure S9B, S9C, S9E, S9F). We noticed, however, that many of these animals did not engorge and instead took smaller blood meals (figure S9D, S9G). Given this, and since neuropeptides are known to act in concert (76), we tested the possibility that sNPF and RYa might act together to regulate blood-feeding behaviour. Indeed, their simultaneous dsRNA injections resulted in about 40% of the animals not taking blood meals (figure 4C). Since RNA interference (RNAi) can be variable, and each female is an independent experiment, we pooled the females that took blood meals and those that didn’t separately before quantifying mRNA from them. We noted that knock-down efficiency was higher in the heads of females that did not take blood meals suggesting that downregulating these genes caused the suppression of blood feeding (figure 4D). Importantly, since these females were offered both blood and sugar, we noted that their sugar feeding was unaffected (figure 4C).

Revised version: Two neuropeptides - sNPF and RYa - showed about 25% and 40% reduced mRNA levels in the heads but the proportion of females that took blood meals remained unchanged (figure S9C, S9D, S9G, S9H). We noticed, however, that many of these animals did not engorge and instead took smaller blood meals (figure S9F, S9J). Given this, and since neuropeptides are known to act in concert (76), we tested the possibility that sNPF and RYa might act together to regulate blood-feeding behaviour. Indeed, their simultaneous dsRNA injections resulted in about 40% of the animals not taking blood meals (figure 4C). Since RNA interference (RNAi) can be variable, and each female is an independent experiment, we pooled the females that took blood meals and those that didn’t separately before quantifying mRNA from them. We noted knockdowns in the heads of females that did not take blood meals suggesting that downregulating these genes caused the suppression of blood feeding (figure 4D). Importantly, since these females were offered both blood and sugar, we noted that their sugar feeding was unaffected (figure 4C).

Earlier version: Since RNAi results in systemic knockdowns and neuropeptides are known to influence behaviour from their action in non-neuronal tissues such as the gut (52, 76, 77), we modified our dsRNA injection protocol to knock sNPF and RYa down in the abdomen while leaving their levels in the head unaltered (figure 4F, 4G, S9J, S9L) (see methods for details). Abdomen-specific knockdown had no effect on blood feeding (figure 4E, S9I, S9K), indicating that abdominal sNPF and RYa alone cannot influence blood feeding and their action in the brain is necessary. In contrast, since every head knockdown of these neuropeptides also abolished their abdominal expression (figure S9C, S9F), we cannot rule out the possibility that both brain- and gut-derived neuropeptides modulate blood feeding. Indeed, since such neuropeptide-mediated gut-brain control of feeding is commonly seen (76), we propose a model, consistent with our data, where brain-derived sNPF/RYamide is essential to trigger blood feeding, while gut-derived peptide acts as a modulatory feedback signal. This remains to be tested with cell-specific tools that are currently not available for this species.

Revised version: Since RNAi results in systemic knockdowns and neuropeptides are known to influence behaviour from their action in non-neuronal tissues such as the gut (52, 76, 77), we modified our dsRNA injection protocol to knock sNPF and RYa down in the abdomen while leaving their levels in the head unaltered (figure 4F, 4G, S9B, S9M, S9N, S9P, S9Q) (see methods for details). Abdomen-specific knockdown had no effect on blood feeding (figure 4E, S9L, S9O), indicating that abdominal sNPF and RYa alone cannot influence blood feeding and their action in the brain is necessary. In contrast, since every head knockdown of these neuropeptides also abolished their abdominal expression (figure S9D, S9E, S9H, S9I), we cannot rule out the possibility that both brain- and gut-derived neuropeptides modulate blood feeding. Indeed, since such neuropeptide-mediated gut-brain control of feeding is commonly seen (76), we propose a model, consistent with our data, where brain-derived sNPF/RYamide is essential to trigger blood feeding, while gut-derived peptide acts as a modulatory feedback signal. This remains to be tested with cell-specific tools that are currently not available for this species.

To determine in which cells sNPF and RYa act, we used hybridisation chain reaction (HCR) (78, 79) in situ. We found that their transcripts are expressed in several non-overlapping neuronal clusters distributed across many regions in the brain (figure 4H, S10A) and only in the Kenyon cells of the mushroom body were they co-expressed (figure 4I, S10B). This is reminiscent of their expression in Drosophila and Bombyx (80–82). Importantly, we found a cluster of cells in the SEZ that expresses the sNPF transcript only in the blood-hungry state (figure 4J, supplementary movie). Its transcripts were either absent or diminished from this cluster in newly emerged or sated females. Since the RNAi experiments suggested the abdominal involvement of these neuropeptides, we also looked for their expression in the midgut where we could detect only sNPF transcripts, and not those of RYa (figure 4K, K’, S11A, S11B, S11E). Interestingly, sNPF transcript expression was also modulated in the midgut enteroendocrine cells (EECs) - hormone-producing cells of the gut. In these cells, sNPF transcripts were specifically increased in blood-hungry females when compared to newly emerged or sated females (figure 4L,M).

Neuropeptides work by acting on their receptors to trigger downstream signalling cascades in the cells that express the receptors. We identified putative receptors for sNPF and RYa (see methods) and found that both were expressed in several overlapping cell clusters in the brain (figure 4N, S10C, S10D). In the midgut, however, we could only detect sNPFR and not RYaR (figure 4N, S11C, S11D, S11F). Since these receptors were identified based on homology, their functional validation will be necessary to implicate them in blood-feeding behaviour.

Together, these data are consistent with a model where increased sNPF expression in both the brain and midgut might drive the blood-feeding behaviour, but that this occurs in the context of RYa’s action in the brain (figure 4O).

Conclusions

Here, we report the blood- and sugar-feeding behaviours of An. stephensi. While sharing key features with other mosquito species, An. stephensi exhibits distinct regulatory patterns. Unlike Ae. aegypti, it does not require mating to initiate blood feeding, though the presence of males is necessary to sustain the enhanced state. In both species, blood feeding is suppressed after a blood meal and until oviposition, but this suppression is mating-dependent in An. stephensi. Since unmated Ae. aegypti females did not blood-feed in our assays, we cannot comment on whether this suppression also requires mating.

Mosquito feeding behaviour is highly context-dependent, shaped by physiological state, environment, and experimental design (36). Our comparisons between Ae. aegypti and An. stephensi were conducted under standardised conditions. While direct comparisons across studies must be made cautiously, patterns observed under particular conditions can offer meaningful insights. For example, virgins of other Anopheles species readily blood-feed (36, 83–85), suggesting that this is typical of the genus, whereas Ae. aegypti virgins do so only under specific conditions and these responses vary among strains (38) (and references within). Whether blood-fed Aedes virgins suppress subsequent feeding is likely also context-dependent with reports suggesting both that they do (86, 87) and do not (42). Post-blood meal suppression of host seeking - well characterised and robust in mated Aedes - appears weaker in mated Anopheles and can be lifted prior to oviposition (46). Indeed, we too observed a stronger post-blood meal suppression of feeding in Ae. aegypti than in An. stephensi (compare Fig. 1C to Fig. 1D). Few studies have investigated this suppression in anopheline virgins, but when they do, they report a weak suppression (88). Taken together, An. stephensi displays characteristics of both culicine and anopheline feeding patterns - its virgin feeding patterns align more closely with other anophelines, yet, like mated Ae. aegypti, it still exhibits clear suppression of blood feeding until oviposition.

What mechanisms might underlie such mating-induced modulation of blood feeding? In Drosophila melanogaster, mating induces a suite of behavioural changes - including increased protein, salt and sugar intake - that are mediated by male-derived sex peptide that acts on central brain circuits in the female brain (4–8). These changes are anticipatory in nature rather than a homeostatic control of feeding behaviour to meet reproductive demands (7, 8). In mosquitoes, the molecular mechanisms driving post-mating behavioural changes are not likely to be the same (89). Instead, mating-derived modulators such as HP1 in Ae. aegypti (59, 90) and male-deposited steroid hormone (20-hydroxyecdysone) in An. gambiae (91, 92) have been shown to induce sexual refractoriness. Whether they also shape feeding decisions remains unknown.

We have identified two neuropeptides - RYa and sNPF - that act synergistically to promote blood feeding in An. stephensi. Our data are consistent with a model where an increase in sNPF levels in the brain - likely released by a cluster of neurons in the SEZ - as well as in the midgut, promotes a state of blood-hunger. Whether this reflects gut-to-brain communication or parallel tissue-local computations to modulate behaviours remains unclear, especially since we detected the transcript of both sNPF and its receptor in both tissues. Inter-organ signalling is well-established in feeding regulation (76, 93). For example, gut-derived NPF in flies can influence brain circuits to increase protein intake (77). More broadly, feeding behaviour is likely regulated by a distributed network of hormonal and neuronal cues across multiple organs. Signals from the ovary, fat body, and other nutrient-sensing tissues may also contribute, acting through both central and peripheral circuits. Many of these inputs remain to be identified.

It is possible that sNPF promotes blood feeding independent of RYa. However, we were unable to knockdown sNPF completely to test this. Its synergistic requirement with RYa, despite no apparent modulation of RYa or its receptor, could result from a general state of hunger that RYa promotes, in the context of which, we were able to uncover the blood-feeding role of the diminished sNPF. Since many neurons co-express these receptors, it is possible that neurons down-stream of sNPF or RYa expressing neurons integrate the two inputs to modulate behaviour.

Earlier version: We note that our findings differ from those reported in Ae. aegypti. In Aedes, sNPF and RYamide act as satiety brakes—peptide or receptor activation curbs host seeking (49, 50, 52) —whereas in Anopheles, our knockdowns show that the same neuropeptides act as a hunger signal, stimulating blood feeding. While these differences are interesting, in the context of their evolutionary histories, and the known differences in how neuropeptides can act differently across species, they are not entirely unexpected. Aedes and Anopheles lineages diverged ≈150-200 million years ago (94), and the two genera have since diverged in almost every facet of biology— diel activity, host preference, feeding and oviposition sites, egg desiccation tolerance, genome size, and chemosensory repertoires among other traits. Neuropeptide pathways are known to drive opposing behaviours in different species: sNPF is known to be orexic in D. melanogaster, A. mellifera, L. decemlineata, B. mori, and P. americana, while it has anorexic effects in S. gregaria (95). Against this backdrop, it is unsurprising that the same neuropeptide pathway can drive opposite behaviours in Aedes and Anopheles. These comparative studies underscore a key lesson for vector control strategies in general. A strategy that is effective in one mosquito species may inadvertently boost vectorial capacity in another. Behavioural interventions, chemical modulators, and even gene drives therefore require rigorous cross-species validation because results from one mosquito species cannot be assumed to hold true for another.

Revised version: We note that our findings differ from those reported in Ae. aegypti. In Aedes, sNPF and RYamide act as satiety brakes—peptide or receptor activation curbs host seeking (49, 50, 52) —whereas in Anopheles, our knockdowns show that the same neuropeptides act as a hunger signal, stimulating blood feeding. While these differences are interesting, in the context of their evolutionary histories, and the known differences in how neuropeptides can act differently across species, they are not entirely unexpected. Aedes and Anopheles lineages diverged ≈150-200 million years ago (94), and the two genera have since diverged in almost every facet of biology— diel activity, host preference, feeding and oviposition sites, egg desiccation tolerance, genome size, and chemosensory repertoires among other traits. Neuropeptide pathways are known to drive opposing behaviours in different species: sNPF is known to be orexic in D. melanogaster, A. mellifera, L. decemlineata, B. mori, and P. americana, while it has anorexic effects in S. gregaria (95). Against this back-drop, it is unsurprising that the same neuropeptide pathway can drive opposite behaviours in Aedes and Anopheles. These comparative studies underscore a key lesson for vector control strategies in general. A strategy that is effective in one mosquito species may inadvertently boost vectorial capacity in another. Behavioural interventions, chemical modulators, and even gene drives therefore require rigorous cross-species validation because results from one mosquito species cannot be assumed to hold true for another.

Methods

Mosquito culture

Anopheles stephensi (Indian strain; Bangalore strain, TIGS-1) and Aedes aegypti (Bangalore collection) were reared in BugDorm cages and maintained at 26-28°C, 70-80% relative humidity with a photo-period of 12h light: 12h dark. All the behavioural assays were performed at these conditions of temperature and humidity. Larvae were fed on a mixture of yeast and dog biscuit (1:3 ratio) and cultured in trays in appropriate densities to prevent over-crowding. Adult mosquitoes were provided ad libitum access to 10% sugar solution (8% sucrose +2% glucose +0.05% methyl paraben +5% vitamin syrup (Polybion SF Complete, Merck Ltd.)) at all times, unless specified otherwise. Females were provided O+ve human blood, secured periodically from blood banks, both for general maintenance and feeding assays.

Regulatory permissions

This project was approved by the Institutional Human Ethics Committee (Ref: inStem/IEC-22/02) and the Institutional Biosafety Committee (TIGS/M-7/8/2020-1).

Earlier version: Blood-feeding assay: first blood meal, 0-120h post-e-mergence (D0-D5)

Revised version: Blood-feeding assay: first blood meal, 0-120h post-emergence (D0-D5)

All An. stephensi feeding experiments were performed from ZT12-ZT15 and Ae. aegypti experiments from ZT10-ZT12 (Zeitgeber time 0 is defined as the time of lights ON). 0-7h-old mosquitoes were cold anaesthetised and collected during the day in 1100ml paper cups covered with a net and secured with a rubber band. These were maintained at the conditions of temperature and humidity described above. To collect virgins, males were separated at the time of collection. 20 females with males (1:1 ratio) or without males were collected for each trial and were aged appropriately with ad libitum access to 10% sugar solution (supplementary figure S1A).

Females were assayed for their blood-feeding behaviour from the day of emergence (D0, 0-7h-old) to 5days of age (D5), every 24h. They were provided constant access to sugar, except when they were offered blood. A detailed schematic has been shown in figure S1A. On the day of the test, appropriately aged females were offered a blood meal via a Hemotek feeder (Hemotek Ltd), which was maintained at 37°C and perfumed with human skin odours. It has been reported that female mosquitoes need to be ‘activated’ to take blood meals. This involves the presentation of CO2 (and other host cues) to the females, likely because CO2 has been shown to gate temperature and visual cues (96, 97). So, prior to offering the blood meal, the mosquitoes were ‘activated’ by the presentation of a hand placed above the cup for 5 minutes, followed by 3 exhalations from the experimenter. Additionally, blood was supplemented with 1mM ATP (Sigma A6419) before loading the feeder (figure 1B). Females were allowed to feed for 60-90mins in dark and abdomens were scored visually for presence of blood, irrespective of the amount ingested (inset, figure S1A). For D0, no sugar was provided at the time of collection. Females were allowed to acclimatise for at least 30 mins in the cups and then offered a choice of both sugar and blood simultaneously, to assess their preference for blood soon after emergence (figure S1A).

In the case of co-housed females, males and females were housed together from the time of collection until the day of the experiment. Males were separated post-hoc and spermathecae of the females were dissected to score for the presence or absence of sperm (figure 1G). All the assays with virgin and co-housed females were performed in parallel by a single experimenter. Mosquitoes from at least 4 independent batches were tested for each day of the experimental timeline. For Ae. aegypti, both virgin and mated females (mating status was confirmed via post-hoc dissection of spermathecae and all co-housed females were found to be mated) were tested for first blood meal at only 8 days post-emergence (D8, figure 1C). Behavioural experiments were performed the same way as described for An. stephensi.

Blood-feeding assay: second blood meal, 24-96h post-first blood meal (D1PBM-D4PBM)

Earlier version: To ensure that the second blood meal is not driven due to an incomplete first blood meal, we selected only fully engorged females to test for subsequent feedings (figure S1A). 5 days-old An. stephensi (both virgins and co-hosued with males) and 7-9 days-old Ae. aegypti (mated, as previously determined) were fed on blood and fully fed females were separated and collected in paper cups as described above. As a female takes about 2 days to digest blood, the second blood meal ingested 24h post-first blood meal (D1PBM) was detected by the presence of Rhodamine B (Sigma R6626), which fluoresces in the red range (figure S1A). Spiking the second blood meal with Rhodamine B (0.04µg/µL of blood) allowed visualisation of the freshly ingested blood, without affecting the feeding efficiency (supplementary figure S1B-D). To assess feeding at subsequent days (D2PBM-D4PBM), abdomens were dissected to differentiate between the bright red freshly ingested blood from the dark and clotted first blood meal. In the case of co-housed An. stephensi females, spermathecae were not dissected to confirm mating as females were presumed to be mated when collected for this assay (5 days-old and co-housed with males for the entire duration from the time of emergence), as previously determined in figure 1G. To confirm whether mating suppresses subsequent feeding (figure 1H), fully fed virgins were collected and allowed to mate for 3-4 days with aged-match males. No males were introduced in the virgin-only controls. Blood-feeding assay was performed as described above. Spermathecae were dissected post-hoc to determine the mating status and unmated females were not considered for the final analysis. Datasets with less than 80% total mating were discarded.

Revised version: To ensure that the second blood meal is not driven due to an incomplete first blood meal, we selected only fully engorged females to test for subsequent feedings (figure S1A). 5 days-old An. stephensi (both virgins and co-housed with males) and 7-9 days-old Ae. aegypti (mated, as previously determined) were fed on blood and fully fed females were separated and collected in paper cups as described above. As a female takes about 2 days to digest blood, the second blood meal ingested 24h post-first blood meal (D1PBM) was detected by the presence of Rhodamine B (Sigma R6626), which fluoresces in the red range (figure S1A). Spiking the second blood meal with Rhodamine B (0.04µg/µL of blood) allowed visualisation of the freshly ingested blood, without affecting the feeding efficiency (supplementary figure S1B-D). To assess feeding at subsequent days (D2PBM-D4PBM), abdomens were dissected to differentiate between the bright red freshly ingested blood from the dark and clotted first blood meal. In the case of co-housed An. stephensi females, spermathecae were not dissected to confirm mating as females were presumed to be mated when collected for this assay (5 days-old and co-housed with males for the entire duration from the time of emergence), as previously determined in figure 1G. To confirm whether mating suppresses subsequent feeding (figure 1H), fully fed virgins were collected and allowed to mate for 3-4 days with aged-match males. No males were introduced in the virgin-only controls. Blood-feeding assay was performed as described above. Spermathecae were dissected post-hoc to determine the mating status and unmated females were not considered for the final analysis. Datasets with less than 80% total mating were discarded.

Blood-feeding assay: post-oviposition, 24-48h post-egg laying (D1PO-D2PO)

For blood-feeding behaviour post-oviposition, fully fed females were separated and a water cup lined with moist paper (ovipositor) was provided for egg laying 2 days post-blood meal (determined in a separate assay described below, supplementary figure S1E; figure S1A). Females were allowed to lay eggs overnight. The following morning, non-gravid females were visually identified by their lean bellies and separated in cups. Blood-feeding assay was performed the same evening for 1 day post-oviposition (D1PO) and the following day (D2PO), as described above. Females were dissected post-hoc and scored for absence of mature ovaries, as a sign of successful oviposition. Gravid females were not considered for the analysis.

Determination of oviposition timing for An. Stephensi

To determine the day of maximum oviposition, 5-6 days-old females co-housed with males were fed on blood, as described above. Fully fed females were separated and collected in BugDorm cages in groups of 10 females/cage. An ovipositor was provided the same day and females were allowed to lay eggs overnight. Eggs laid were counted the next day and a fresh ovipositor was provided. This cycle was repeated for the next 5 days. Total eggs laid each day were divided by the number of females present at the time of counting (supplementary figure S1E).

Sugar-feeding assay for An. Stephensi

As mosquitoes feed on sugar in small bouts, as opposed to a replete blood meal, we assessed sugar-feeding behaviour of females in windows of 24h for each test day. 0-7h-old mosquitoes were collected during the day, separated in cups and aged appropriately. They were given access to 10% sugar solution at all times, as described above, except for the day of emergence (D0) when no sugar was provided at the time of collection (similar to the blood-feeding assay performed at D0).

To assess their sugar-feeding behaviour on the day of emergence (D0), 0-7h-old females were acclimatised in the cups for 1h and sugar spiked with Rhodamine B (1mg/25ml of sugar; Sigma R6626) was presented for 2.5-3h. The fluorescence of Rhodamine B in the red channel allowed visualisation of even the tiniest amounts of sugar ingested by a female (figure S1G). For D1 feeding, 3-10h-old females (with access to normal sugar so far) were then provided with Rhodamine B-sugar and allowed to feed on it for 24h. For D2 feeding, 27-34h-old females (with access to normal sugar so far) were fed on Rhodamine B-sugar for the next 24h and scored as D2. This cycle was repeated until D5. In the case of co-housed females, males were separated post-hoc and spermathecae of the females were dissected to score for the presence or absence of sperm (figure S1H).

To assess sugar feeding post-blood meal, females fully fed on blood were collected and separated in cups, as described above. They were allowed to feed on Rhodamine B-sugar for 24h following a blood meal and scored as D1PBM. The next set of females (with access to normal sugar so far) were then provided with Rhodamine B-sugar from 24h-48h post-blood meal and scored as D2PBM. This cycle was repeated until 4 days post-blood meal (D4PBM, figure 1E).

For sugar feeding post-oviposition, blood-fed females were provided an ovipositor 3 days later and were forced to oviposit during daytime by artificially creating dark conditions. Oviposited females (non-gravid) were then visually identified, separated and presented with Rhodamine B-sugar for the following 24h and scored as D1PO. The next set (with access to normal sugar so far) was then provided with Rhodamine B-sugar, 24-48h post-oviposition and scored as D2PO(figure 1E). For each day tested, females were dissected post-hoc and scored for absence of mature ovaries, as a sign of successful oviposition. Gravid females were not considered for the analysis.

Dual choice assay of blood and sugar

Earlier version: Mosquitoes were collected as described above. Males and females were co-housed for 4-6 days with ad libitum access to sugar. To assess the choice between sugar and blood, mosquitoes were either given continuous access to normal sugar (sugar-sated) or starved with water 24h prior to the assay (sugar-starved). At the time of the assay, post-activation (as described above), blood and Rhodamine B-spiked sugar (1mg/25ml of sugar, Sigma R6626) were provided in parallel. After 60mins, abdomens were scored visually for ingestion of blood under white light, and under fluorescent microscope (red channel) for ingestion of sugar (figure S1F). Co-housed males served as an internal control for sugar feeding. As males are unable to feed on blood, they were only scored for presence or absence of Rhodamine B-spiked sugar. Results are presented in figure 1F.

Revised version: Mosquitoes were collected as described above. Males and females were cohoused for 4-6 days with ad libitum access to sugar. To assess the choice between sugar and blood, mosquitoes were either given continuous access to normal sugar (sugar-sated) or starved with water 24h prior to the assay (sugar-starved). At the time of the assay, post-activation (as described above), blood and Rhodamine B-spiked sugar (1mg/25ml of sugar, Sigma R6626) were provided in parallel. After 60mins, abdomens were scored visually for ingestion of blood under white light, and under fluorescent microscope (red channel) for ingestion of sugar (figure S1F). Co-housed males served as an internal control for sugar feeding. As males are unable to feed on blood, they were only scored for presence or absence of Rhodamine B-spiked sugar. Results are presented in figure 1F.

Statistical analysis of feeding behaviours

Statistical analysis was performed using R (version 4.4.2) and Graph-Pad Prism (v10.2.0). Blood- and sugar-feeding behaviours of virgins and co-housed females in figure 1D, E were analysed using Generalized Linear Mixed-effects Models (GLMMs) with a binomial distribution and logit link function. The models were fitted using the “glmer” function from the “lme4” package in R. The probabilities of blood- or sugar-feeding, both pre- and post-blood meal were modeled as a function of Day (age) and Group (virgin vs. co-housed with males), including their interaction term to test whether group differences varied across days. Replicate identities (BioID) and larval batches (only for first blood meal in sugar-fed females) were included as random factors to account for replicate-level and batch-level variations respectively. Post-hoc analyses were conducted using estimated marginal means (emmeans package) for age-effects within groups (p-value adjusted for multiple comparisons using Tukey’s method) and pairwise comparison between groups at each day (p-value adjusted using Bonferroni method). To test whether mating success altered the sugar- and blood-feeding propensity of females co-housed with males, the probabilities of feeding were modelled only the co-housed data, as a function of Day (age) and Mating (yes or no), including their interaction term and accounting for variation due to replicate IDs (BioID). Statistical significance was set at α = 0.05 for all analyses. Complete model outputs are provided in supplementary sheet 1.

Statistical analyses for data in Figures 1C, 1F, and 1H were performed using GraphPad Prism (v10.2.0) with appropriate tests as indicated in figure legends.

Y-maze olfactory assay for host-seeking behaviour in An. Stephensi

A 30inch Y-maze was constructed as per the WHO guidelines (https://www.who.int/publications/i/item/9789241505024). Each arm of the Y-maze had a trapping port, while the stem was attached to a holding port. Each port had a mesh screen at one end, while the other was fitted with a sliding door to trap or release mosquitoes. The holding port was attached to a fan which pulled air from the two arms, over the mosquitoes, and out at a velocity of 0.3-0.6m/s. Optic flow information (alternate black and white stripes) was placed underneath the Y-maze to act as visual guide for navigation.

Earlier version: Following conditions were tested: D0 virgins, D5 sugar-fed virgins, D5 sugar-fed mated females (D5(m)), 3 days post-blood meal virgins (D3PBM), 3 days post-blood meal mated females (D3PBM(m)), and 2 days post-oviposition mated females (D2PO(m)). For mated conditions, females were co-housed with males for 5 days post-emergence and assumed to be mated based on our previous findings (figure 1G). For each condition, mosquitoes were collected and provided ad libitum access to sugar as described above. Females were transferred to holding ports 7-8h prior to the assay and were starved during this time with only access to water. For each run, mosquitoes were acclimatised for 5 minutes: 2 min with clean air and 3 minutes with an experimenter presenting a hand at one of the arms and simultaneously exhaling at resting rates. To account for disruption in airflow, the hand of a mannequin was placed at the other arm (control). Postacclimatisation, the door of the holding port was opened to release the mosquitoes and they were allowed to make a choice for 5 minutes in the presence of continued host cues. The ports were then closed and mosquitoes in each arm counted. All the mosquitoes trapped in a port, sitting on the door or resting on the half-length of the Y-arm were considered as “attracted” to that stimuli (host or control, figure 2A). Participation percentage was also calculated as number of mosquitoes that passed half-way through the Y-stem/total number of mosquitoes released (figure S1I).

Revised version: Following conditions were tested: D0 virgins, D5 sugar-fed virgins, D5 sugar-fed mated females (D5(m)), 3 days post-blood meal virgins (D3PBM), 3 days post-blood meal mated females (D3PBM(m)), and 2 days post-oviposition mated females (D2PO(m)). For mated conditions, females were co-housed with males for 5 days post-emergence and assumed to be mated based on our previous findings (figure 1G). For each condition, mosquitoes were collected and provided ad libitum access to sugar as described above. Females were transferred to holding ports 7-8h prior to the assay and were starved during this time with only access to water. For each run, mosquitoes were acclimatised for 5 minutes: 2 min with clean air and 3 minutes with an experimenter presenting a hand at one of the arms and simultaneously exhaling at resting rates. To account for disruption in airflow, the hand of a mannequin was placed at the other arm (control). Post-acclimatisation, the door of the holding port was opened to release the mosquitoes and they were allowed to make a choice for 5 minutes in the presence of continued host cues. The ports were then closed and mosquitoes in each arm counted. All the mosquitoes trapped in a port, sitting on the door or resting on the half-length of the Y-arm were considered as “attracted” to that stimuli (host or control, figure 2A). Participation percentage was also calculated as number of mosquitoes that passed half-way through the Y-stem/total number of mosquitoes released (figure S1I).

Several controls were included to avoid experimental artefacts. Each experimental day started with a blank trial (where no stimuli were presented on either arm of the Y-maze) and the side of the host cues presentation was alternated with each trial to eliminate side biases. Trials for blood-fed virgins and mated females were done both in parallel and on multiple days to eliminate the influence of external factors which might be present on a given day. When testing blood-fed females, age-matched sugar-fed females (blood-hungry) were included wherever possible as positive controls. These females consistently showed attraction to host cues, as expected. Lastly, we made sure that at least 50% of the mosquitoes had made either choice in all runs, except for D0 (newly emerged females) and blank runs. In the rare cases that this was not true, we did not consider those for the analysis. Statistical analysis was performed using GraphPad Prism (v10.2.0). To compare the attraction of females to host cues and control in figure 2B, each time-point was treated as an independent experiment and p-values were computed for each pair-wise comparison.

Bulk RNA sequencing of central brains and data analysis

For bulk RNA sequencing, central brains (ie. brains without optic lobes) were dissected from the following conditions: D0 and D5 sugar-fed males; D0 and D5 sugar-fed virgin females; 1 day post-blood meal virgin (D1PBM) and mated females (D1PBM(m)) and 1 day post-oviposition mated females (D1PO(m)). For post-blood meal conditions, 5 days-old females were fed on blood and fully fed females were separated for sample preparation. Females co-housed with males for 5 days were assumed to be mated based on our previous findings (figure 1G). Three replicates from independent batches were prepared for each condition.

Mosquitoes were cold anaesthetised before dissections. Central brains (with optic lobes removed) were dissected in RNase-free 1X PBS and immediately placed into ACME solution (13:3:2:2 of water, methanol, acetic acid, and glycerol) on ice, which mildly fixes the tissues (98). For each condition, 100 brains were dissected per replicate. The mildly fixed brains were washed with cold 1X PBS and centrifuged at 13,000 rpm for 5 min at 4°C to remove ACME. Brains were then homogenised in TRIzol (Invitrogen), flash-frozen, and stored at -80°C until RNA extraction. RNA was extracted using Direct-zol RNA MicroPrep (Zymo Research) columns. RNA quantity and quality were assessed using the ThermoFisher Qubit 4 and RNA 6000 Nano kit Bioanalyser, respectively.

The cDNA library was prepared for each sample using the NEBNext Ultra II Directional RNA Library Prep kit (NEB E7765L), with an average fragment size of 250 bp. Some of the library samples were subjected to size and quality checks using a bioanalyser. Paired-end sequencing was carried out on the Illumina Hiseq 2500 sequencing platform, generating 2×100 bp reads. The library preparation and sequencing were performed at the Next Generation Genomics Facility (NGGF) at Bangalore LifeScience Cluster (BLiSC), India. Post-sequencing, samples were demultiplexed and delivered as raw Fastq files. After passing the quality check using the FastQC tool, the reads were mapped onto the reference An. stephensi transcriptome (NCBI RefSeq assembly GCF013141755.1) (66) using Kallisto v0.43.1 (99) with default parameters. 81-87% of reads aligned to the mosquito transcriptome.

Transcript-level abundance was quantified using Kallisto and gene-level summarisation was performed using the Tximport function in R, with transcript isoform information retained (supplementary sheet 2). Using the EgdeR package, the raw counts were converted to counts per million (CPM) and log2 transformed before data filtering (genes with less than 1 CPM in at least 3 or more samples were filtered out) and normalisation, using the trimmed mean of M-values (TMM) method. Principal component analysis (PCA) was performed using the prcomp function in R and visualised using the ggplot function. Data was checked for batch effects using gPCA (100) in R and the test statistic “delta” was found to be 6.957686e-05, with a p-value of <0.001, indicating no batch-effects.

Differential gene expression analysis

Differential expression analysis was performed using the Limma package in R. A design matrix was created to specify the sample conditions, and the Voom function was used to estimate the mean-variance relationship across genes. A linear model was then fitted to the data using the mean-variance relationship. Bayesian statistics were calculated for the data by creating a contrast matrix to extract the linear model fit. To identify differentially expressed genes (DEGs), the “decide-Tests” function was used by applying a fold change threshold of 2 (log2fc=1) and a FDR (false discovery rate) adjusted p-value cut-off of 0.05. The number of genes identified to be upregulated and downregulated in each pair-wise comparison are shown in figure S2A and the gene lists are provided in supplementary sheet 3.

To plot transcripts per million (TPM) values as heatmaps, the Fastq files for the central brain samples were mapped onto the reference An. stephensi transcriptome and abundance was quantified using Kallisto the same way as described above. Gene-level summarisation was done using Tx-import as described above, except all the isoforms of a given transcript were mapped to a single gene (supplementary sheet 2). The resultant TPM values were first filtered for TPM>0.1 in each sample. Next, the genes that were not quantified in at least 2 of the 3 replicates of each sample group were removed. The filtered dataset were then log10 transformed and heatmaps were plotted using the pheatmap function in R. To construct heatmaps of genes involved in carbohydrate and lipid metabolism (supplementary figure S3, S4), key enzymes of the respective canonical pathways were manually curated based on the existing literature. Gene annotations for neuropeptides and neuropeptide receptors were curated from published sources (101–103). Fat body, midgut and ovary data (SRX620223, SRX618937 and SRX620224; bioproject SRP043489) for 2-4 days-old, sugar-fed An. stephensi females were taken from Prasad et al.(2017) (104). Fastq files were downloaded from the SRA explorer and processed similarly (supplementary sheet 2).

Gene set enrichment analysis

eggNOG mapper tool 2.1.9 (105) (parameters: 70% identity and minimum 70% query coverage) was used to find functionally annotated orthologs of the identified DEGs in An. gambiae. An. gambiae-specific orthologs were fed into the GOSt tool from gProfiler2 (106) toolset for functional enrichment analysis. A custom gene list comprising of all the genes identified in our RNA-seq experiment was provided as a background instead of the whole genome. Using this background, an ordered query was run with an FDR adjusted p-value threshold of 0.05 to determine which GO terms (Biological Processes) were over-represented. This was done for all sets of DEGs and the top 10 terms of each set were plotted using ggplot function in R (figure S2B).

dsRNA-mediated gene knockdown

To synthesise the double-stranded RNA (dsRNA), coding regions (500-600bp) of candidate genes were amplified from cDNA of An. stephensi females using forward and reverse primers with T7 promoter sequence at 5’ ends. The gel purified fragments were then used as a template for in vitro transcription reactions using Megascript RNAi kit (Invitrogen) or High-Yield T7 RNAi Kit (Jena Biosciences) as per the manufacturers’ protocols. The dsRNAs were purified using the phenol-chloroform method and pellet dissolved in nuclease-free water. The dsRNA for green fluorescent protein (GFP) was synthesised from plasmid pAC5.1B-EGFP (Addgene 21181) to be used as a control for injection-related behavioural changes.

Earlier version: Virgins were collected for injections as described above. The day and amount of dsRNA to be injected were determined for each target gene individually. dsEiger: 2.5-3.6µg injected per female in 24h-old virgins; dsHairy: 2-3µg injected per female in 24h-old virgins; dsJHEH: 2.5-3µg injected per female in 72h-old virgins; dsPrp: 1.5-3µg injected per female in 72h-old virgins; dsOrk: 2.5-3µg injected per female in 96h-old virgins; ds118504077 and ds118504075: 3-3.6µg injected per female in 96h-old virgins. For dual injections of ds118504077+ ds118504075, 3.5-4.2µg of a 1:1 mixture of purified dsRNAs was injected per female in 96h-old virgins.

Revised version: Virgins were collected for injections as described above. The day and amount of dsRNA to be injected were determined for each target gene individually. dsEiger: 2.5-3.6µg injected per female in 24h-old virgins; dsHairy: 2-4.3µg injected per female in 24h-old virgins; dsJHEH: 2.5-3µg injected per female in 72h-old virgins; dsPrp: 1.5-3µg injected per female in 72h-old virgins; dsOrk: 2.5-3µg injected per female in 96h-old virgins; ds118504077 and ds118504075: 3-3.6µg injected per female in 96h-old virgins. For dual injections of ds118504077+ ds118504075, 3.5-4.2µg of a 1:1 mixture of purified dsRNAs was injected per female in 96h-old virgins.

Earlier version: To achieve tissue-specific knockdown (abdomen only or head+abdomen) of sNPF and RYa, injections were performed at different concentrations and on different days. Injecting dsRNA into 0-10h-old females produced abdomen-specific knockdowns without affecting head expression, whereas injections into 96h-old females resulted in knockdowns in both tissues. Moreover, head knockdowns in older females (96h-old) required higher dsRNA concentrations (3.5-4.2µg injected per female), with knockdown efficiency cor-relating with the amount injected. In contrast, abdominal knockdowns in younger females (0-10h-old) could be achieved even with lower dsRNA amounts (2-2.4µg or 3.5-4.2µg injected per female). For dual injections, 3.5-4.2µg of the 1:1 mixture of purified dsRNAs was injected per female as described above: in 0-10h-old females for abdomen-specific knockdowns and in 96h-old virgins for head+abdomen knockdowns.

Revised version: To achieve tissue-specific knockdown (abdomen only or head+abdomen) of sNPF and RYa, injections were performed at different concentrations and on different days. Injecting dsRNA into the abdomens of 3-7h-old females produced abdomen-specific knockdowns without affecting head expression, whereas injections into head or abdomens of 96h-old females resulted in knockdowns in both tissues. Moreover, head knockdowns in older females (96h-old) required higher dsRNA concentrations (3.5-4.2µg injected per female), with knockdown efficiency correlating with the amount injected. In contrast, abdominal knockdowns in younger females (3-7h-old) could be achieved even with lower dsRNA amounts (2-4.2µg injected per female). For dual injections, 3.5-4.2µg of the 1:1 mixture of purified dsR-NAs was injected per female as described above: in 3-7h-old females for abdomen-specific knockdowns and in 96h-old virgins for head+abdomen knockdowns (figure S9B).

Injections were performed on a FemtoJet 4i system (Eppendorf) and each female was injected with 0.5-0.6µL of the respective concentrated dsRNAs in the upper thorax. In some cases, dsRNAs were also injected in the abdomens to ensure that the site of injection does not alter the behaviour. Uninjected or dsGFP-injected (same concentration as that of the candidate gene) females were included as a control each time. 15-25 females were injected for each trial and at least 2 trials from independent batches of mosquitoes were performed for each target gene.

Earlier version: Injected females were allowed to recover in paper cups with ad libitum access to sugar and assessed for blood-feeding behaviour at 3 days post-injection, as previously described. Fed and unfed females were visually scored and 5-6 females were randomly selected for each sample-whole heads and abdomens were collected separately for analysis. Total RNA was extracted using Direct-zol RNA MicroPrep (Zymo Research) columns, as per the manufacturer’s protocol. RNA was quantified using Qubit Broad range assay (Invitrogen). For each sample, 200-220ng RNA was used to prepare cDNA using the SuperScript™ IV Reverse Transcriptase (Invitrogen 18090050) or Maxima H Minus Reverse transcriptase (Thermo Scientific EP0752). Quantitative PCR (qPCR) was performed using PowerUp™ SYBR™ Green Master Mix (Applied Biosystems) to determine the knockdown efficiency. Relative mRNA abundances were calculated using the delta-delta Ct method (107), normalised to RpS7. Statistical analyses were performed using GraphPad Prism (v10.2.0). All the primers used in the study are listed in supplementary table 1.

Revised version: Injected females were allowed to recover in paper cups with ad libitum access to sugar and assessed for blood-feeding behaviour at 3 days post-injection, as previously described. Fed and unfed females were visually scored and 5-6 females were randomly selected for each sample-whole heads and abdomens were collected separately for analysis. Total RNA was extracted using Direct-zol RNA MicroPrep (Zymo Research) columns, as per the manufacturer’s protocol. RNA was quantified using Qubit Broad range assay (Invitrogen). For each sample, 200-220ng RNA was used to prepare cDNA using the SuperScript™ IV Reverse Transcriptase (Invitrogen 18090050) or Maxima H Minus Reverse transcriptase (Thermo Scientific EP0752). Quantitative PCR (qPCR) was performed using PowerUp™ SYBR™ Green Master Mix (Applied Biosystems) to determine the knock-down efficiency. Relative mRNA abundances were calculated using the delta-delta Ct method (107), normalised to RpS7. Statistical analyses were performed using GraphPad Prism (v10.2.0). All the primers used in the study are listed in supplementary table 1.

To assess the sugar-feeding behaviour of dsRNA injected females (dssNPF, dsRYa and dsRYa+sNPF), injections were performed as described above. Injected females were allowed to recover with ad libitum access to normal sugar for the first two days, which was then replaced with Rhodamine B-spiked sugar (1mg/25ml of sugar, Sigma R6626) for the next 24h. Control females were treated similarly. The blood-feeding assay was performed 3 days post-injection as described above, with both blood and sugar provided in parallel. Abdomens were scored visually for ingestion of blood under white light, and under fluorescent microscope (red channel) for ingestion of sugar (figure S9A).

Hybridisation chain reaction (HCR)

Previously published protocol (108) was modified to perform HCR in situ hybridisations on brain and gut tissues collected from females under three different conditions: D0 (not interested in blood), D5 virgins (blood-hungry) and mated D2PBM(m) (blood-sated). All the reagents, including the probes, amplifiers and buffers were procured from Molecular Instruments Inc. A set of custom-designed 20 probe pairs were used per target gene: sNPF (118511038), RYa (118513478) and their receptors. To find sNPF and RYa receptors in An. stephensi, putative protein sequences were blasted against orthologs in Drosophila melanogaster, Ae. aegypti and An. gambiae. Based on high identities, consistent in all three organisms, 118517381 and 118507167 (NPY receptor type 2) were identified as RYa receptor (RYaR) and sNPF receptor (sNPFR), respectively.

5-7 brains for each target, per condition were dissected in ice-cold PBS and fixed with 4% PFA (in PBS) for 20 minutes at room temperature. Post-washing steps with PTw (0.1% Tween-20 in PBS), 4 times 10 mins each, tissues were permeabilised using detergent solution (50mM Tris-Cl, pH7.5, 150mM NaCl, 1mM EDTA, pH8, 0.5% Tween-20, 1% SDS) for 30 mins at room temperature. Brains were then prehybridised in probe hybridisation buffer for 30 mins at 37°C, followed by incubation with neuropeptide probes (0.8pmol each) at 37°C for 48h in a humid chamber. Samples were washed with pre-warmed wash buffer, 4 times 15 mins each, at 37°C to remove unbound probes. This was followed by 2 washes, 5 mins each with 5X SSCT (0.1% Tween-20 in 5X SSC) at room temperature. Samples were then pre-amplified using amplification buffer for 30 mins at room temperature. For each sample, 2µL (3µM stock) of each hairpins were mixed and prepared by heating at 95°C for 90secs for annealing. They were allowed to cool to room temperature for 30mins in a dark chamber and 100µL amplification buffer was added to the mix. Brains were incubated in the hairpin solution for 12-16h at 37°C in dark and excess hairpins were washed off with multiple 5X SSCT washes: twice for 5 mins, twice for 30mins, once for 5 mins. For nc82 staining, all the steps were performed in 0.5% PTw buffer (0.5% Tween-20 in PBS). Brains were equilibrated in 0.5% PTw and incubated with anti-mouse nc82 (1:200; DSHB, RRID: AB2314866) for three overnights at 4°C. Post-washing, goat anti-mouse Alexa-fluor405 (1:400, Invitrogen) was added and samples were left in dark, for one-two overnights at 4°C. After washing off the unbound secondary antibody, samples were equilibrated for at least an hour in Vectashield and mounted in Vectashield for imaging. In situs against neuropeptide receptors were performed the same way, except the washing steps to remove the unbound probes were performed at 42°C and samples were incubated with the hairpin solution at room temperature (to minimise background). As controls for background signal, samples were prepared the same way, except no probes were added in the hybridisation buffer (figure S10E). Samples for all three conditions were dissected, processed and imaged the same day.

Protocol for gut in situs was adapted from Slaidina et al. (109). 5-8 gut samples for each target, per condition were dissected in ice-cold PBS and fixed with 4% PFA (in PTw) for 20 minutes at room temperature. Post-washing steps with PTw, tissues were dehydrated on ice with graded methanol washes: 25% methanol (in PTw), 50% methanol (in PTw), 75% methanol (in PTw) and 100% methanol for 10 mins each. Samples were stored in 100% methanol for up to a month, at -20°C. For rehydration, samples were washed sequentially with graded methanol: 75% methanol (in PTw), 50% methanol (in PTw), 25% methanol (in PTw), followed by a wash in permeabilisation solution (1% Triton-X in PBS). Samples were permeabilised for 2h at room temperature and fixed again with 4% PFA (in PTw) for 20 minutes. Following 10mins washes: washed once with PTw, twice with 50% PTw/50% 5X SSCT and twice with 5X SSCT, samples were pre-hybridised in probe hybridisation buffer for 30 mins at 37°C, followed by incubation with probes (0.8pmol each) at 37°C for 48h in a humid chamber. Post washing steps with wash buffer at 37°C and 5X SSCT at room temperature, samples were pre-amplified in amplification buffer for 30 mins at room temperature. Hairpins were prepared the same way as described above and samples were incubated in hairpin solutions for 12-16h in dark at room temperature. Following 2 washes 5 mins each with 5X SSCT, samples were incubated with DAPI (1:1000 in 5X SSCT) for one hour at room temperature. Samples were washed with 5X SSCT, 4 times, 10mins each and equilibrated in Vectashield for at least an hour before mounting. As controls for background signal, samples were prepared the same way, except no probes were added in the hybridisation buffer (figure S11G, H). Samples for different conditions, prepared on different days and stored at -20°C, were processed and imaged together. Statistical analyses were performed using GraphPad Prism (v10.2.0).

All samples were imaged at Central Imaging and Flow Cytometry Facility (CIFF) at the Bangalore Life Science Cluster (BLiSC), India. Images were acquired with a 40X/1.3NA or 10X/0.4NA (whole guts) objectives at a resolution of 512X512 pixels, on an Olympus FV3000 system. Images were viewed and manually analysed using ImageJ/FIJI (110). Maximum intensity projections of all z-stacks in an image are shown, except in figures 4H, 4I, S10B, S10D, where stacks were manually selected to capture relevant brain structures.